```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_midterm_exam_q", "week_6", "midterm", assignment_points = 110.0, assignment_tag = 'week6-midterm')

# Initialize Otter
import otter
grader = otter.Notebook("1_midterm_exam_q.ipynb")
```

# ❓ ENGR131: Introduction to Programming for Engineers

You will have one hour to complete the midterm exam. Please read the instructions carefully. You will receive feedback similar to your practice midterm exam.

## Recommended Tests Taking Strategies

1. Look over the point totals for each question
2. Quickly complete the multiple choice questions
3. Move on to the short answer questions and complete the ones you know first.
4. Go back and refine your answers to the short answer questions -- this is a good opportunity to use the python interpreter to test your code.

For the multiple choice questions it is expected that you have access to a python interpreter. Do not go at it alone, use the resources available to you to help you answer the questions. Without the interpreter some of the questions may be difficult to answer.

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_midterm_exam_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_midterm_exam_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_midterm_exam_q import Question3
Question3().show()

```

## Question 1 (Points: 16.0): Quality Control in Manufacturing: Loops, Continue, Break, Dictionaries, and Lists

### Instructions:

You are managing a **quality control system** in a manufacturing plant. Each manufactured part is represented as a dictionary, where the keys are **part IDs** (e.g., `"P1"`) and the values contain details about the part. Your task is to implement a function `inspect_parts()` that processes a list of parts using **loops**, `continue`, and `break`.

### Requirements:

1. **Input:** The function receives a list of dictionaries, each representing a part:

| id | status   | weight (kg) |
|----|----------|------------|
| P1 | passed   | 1.2        |
| P2 | passed   | 2.3        |
| P3 | rejected | 3.8        |
| P4 | passed   | 5.1        |

2. **Part 1:** Count the number of parts with status `"passed"` and return the count.
3. **Part 2:** Skip parts labeled as `"rejected"`, and create a list of IDs for accepted parts.
4. **Part 3:** If a part's weight exceeds `4.5 kg`, **break the loop** and return `"Overweight part detected!"`.
5. **Part 4:** Return a dictionary summarizing:
- The count of `"passed"` parts.
- The list of valid part IDs.
- The break message (if triggered).

### Helpful Information:

- Use a loop to iterate through the `parts` list.
- Use `continue` to **skip** certain iterations based on a condition.
- Use `break` to **exit** the loop early when another condition is met.
- Access dictionary values using `dictionary[key]`.
- Useful methods:
    - `.append()`: Adds an element to a list.

Follow the provided step-by-step instructions carefully and complete each part.

```python
# Define a list of dictionaries representing parts in the quality control system
...

# Define the function `inspect_parts` to process parts
...

    # Initialize variables
    # Create a variable `approved_count` and set it to 0
    # Create a variable `accepted_parts` and set it to an empty list
    # Create a variable `alert_message` and set it to None
    ...
    
    # Loop through parts
    ...
    
        # Part 3: Check for overweight parts
        # If the part's weight is greater than 4.5, set `alert_message` to "Overweight part detected!" and break the loop
        ...
        
        # Part 2: Skip parts with status "rejected"
        # If the part's status is "rejected", continue to the next iteration
        ...
        
        # Part 1: Count parts with status "passed"
        # If the part's status is "passed", increment `approved_count` by 1 and append the part's id to `accepted_parts`
        ...
        
    # Part 4: Return the results as a dictionary
    # Return a dictionary with keys "approved_count", "accepted_parts", and "alert_message" and their corresponding values
    ...
    
# Call the function with the parts dictionary and store the result in a variable `qc_result`
...
```

```python
grader.check("loops-and-dictionaries-quality-control")
```

## Question 2 (Points: 13.0): Debugging Python Functions: Fixing Syntax, Runtime, and Logical Errors
### Instructions:

You are provided with the following **buggy code** related to **thermodynamics**. Your task is to debug the function `convert_kelvin_to_celsius`, which is intended to take a list of temperatures in Kelvin and return a new list where only the **valid (above absolute zero)** temperatures are converted to Celsius using the formula:

$$
C = K - 273.15
$$

```python
def convert_kelvin_to_celsius(temps)

celsius_values = []
for t in temp:
    if t < 273.15
    contine

    celsiuss_values.add(t - 273.5
    print(celsius_values)

# Example usage:
temperatures = [300, 250, 280, 260, 310]
print(convert_kelvin_to_celsius(temperatures)) # Expected output: [26.85, 6.85, 36.85]
```
### Requirements:

1. **Input:** The function receives a list of temperatures in **Kelvin**.
2. **Expected Behavior:** The function should return a new list where **only valid Kelvin temperatures** (≥ 273.15) are converted to Celsius.
3. **Possible Errors to Fix:**
- **Syntax Errors** (missing colons, incorrect operators, typos).
- **Runtime Errors** (undefined variables, incorrect function calls).
- **Logical Errors** (incorrect formulas or conditions).
4. **Return Value:** The function should **return** the converted list, not just print it.

```python
# Define a function `convert_kelvin_to_celsius` that takes a list of Kelvin temperatures as input
...
```

```python
# Do Not Modify Below This Line
temperatures = [300, 250, 280, 260, 310]
result = convert_kelvin_to_celsius(temperatures)
print(result)
```

```python
grader.check("loops-and-errors-debugging-thermodynamics")
```

## Question 3 (Points: 16.0): Numerical Computing with NumPy: Positive Multiple of 7 Checker 🔢
### Instructions:

You are working with numerical data and need to determine which numbers in a given NumPy array are **positive multiples of 7**. Your task is to write a function `find_multiples_of_7()` that takes a NumPy array of integers and returns a boolean array indicating which elements are **positive multiples of 7**.

### Requirements:

1. The function receives a NumPy array of integers, or a list of integers. Hint - You can convert a list to a NumPy array using `np.array()`.
2. Implement an efficient multiple-checking function using NumPy.
3. The function should return a boolean NumPy array of the same shape as the input, where `True` represents a **positive** multiple of 7, and `False` otherwise.

### Helpful Information:

- A multiple of **7** is any number `n` that is evenly divisible by **7**.
- The first few positive multiples of 7 are: **7, 14, 21, 28, 35, ...**
- The function should **only** return `True` for **positive multiples of 7** (i.e., negative numbers should return `False`, even if they are divisible by 7).

**Make sure to include docstrings and comments in your code to explain the logic and purpose of the function.**

```python
# Import NumPy using the common alias np
...

# Define the function `find_multiples_of_7` which accepts an input array, or list
...
    """
    ...
    """
    # Ensure input is a NumPy array
    ...

    # Create a boolean array where `True` represents positive multiples of 7
    ...

    # Return the boolean array `is_multiple_of_7`
    ...

# Example usage:
test_array = np.array([0, 7, 14, -7, 21, -14, 28, 35, -49, 50])
multiples_results = find_multiples_of_7(test_array)
print(multiples_results)
```

```python
grader.check("numpy-multiple-of-7-check")
```

## Question 4 (Points: 10.0): Bonus In-Place Removal of Duplicates in a Sorted Array 🔢
### Instructions:

You are given an integer array `nums` sorted in **non-decreasing order**. Your task is to **remove duplicates in-place** such that each unique element appears **only once**. The **relative order** of the elements should be maintained.

### Requirements:

1. **Modify `nums` in-place** such that the **first `k` elements** contain only **unique elements**, while the rest of the elements do not matter.
2. **Return `k`**, the number of unique elements in `nums`.
3. The function **should not** use extra space; it must modify `nums` in-place.
4. the function must return 0 if the input list is empty.

### Example Usage:
```python
nums = [1, 1, 2, 2, 3, 4, 4]
k = remove_duplicates(nums)
print(k, nums[:k])  # Expected output: 4, [1, 2, 3, 4]
```

**Constraints:**
- `1 <= len(nums) <= 10^4`
- `-10^4 <= nums[i] <= 10^4`
- `nums` is sorted in **non-decreasing order**.

```python
# Define the function `remove_duplicates` which modifies `nums` in-place
...

# Example usage:
nums = [1, 1, 2, 2, 3, 4, 4]
k = remove_duplicates(nums)
print(k, nums[:k])  # Expected output: 4, [1, 2, 3, 4]
```

```python
grader.check("remove-duplicates-sorted-array")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week6-midterm", "1_midterm_exam_q")
```

Tests for loops-and-dictionaries-quality-control (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(16.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(55.0)
>>> log_variable('total-points', f'week6-midterm, 1_midterm_exam_q', 55.0)
>>> question_id = 'loops-and-dictionaries-quality-control-1'
>>> max_score = 1.0
>>> score = 0
>>> function_exists = callable(inspect_parts)
>>> assert function_exists, 'The function `inspect_parts` must be defined and callable.'
>>> if (function_exists, 'The function `inspect_parts` must be defined and callable.'):
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(function_exists))
```
failure_message: Failed: The function `inspect_parts` is not defined correctly.
success_message: Success: The function `inspect_parts` exists and is callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import inspect
>>> import re
>>> question_id = 'loops-and-dictionaries-quality-control-2'
>>> max_score = 2.0
>>> score = 0
>>> code = inspect.getsource(inspect_parts)
>>> approved_count_init = re.search('approved_count\\s*=\\s*0', code)
>>> accepted_parts_init = re.search('accepted_parts\\s*=\\s*\\[\\]', code)
>>> alert_message_init = re.search('alert_message\\s*=\\s*None', code)
>>> initial_values = all([approved_count_init, accepted_parts_init, alert_message_init])
>>> assert initial_values, 'Expected variables `approved_count = 0`, `accepted_parts = []`, and `alert_message = None` in the function.'
>>> if (initial_values, 'Expected variables `approved_count = 0`, `accepted_parts = []`, and `alert_message = None` in the function.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(initial_values))
```
failure_message: Failed: The function does not correctly initialize variables.
success_message: Success: The function correctly initializes variables!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-quality-control-3'
>>> max_score = 3.0
>>> score = 0
>>> test_parts = [{'id': 'X1', 'status': 'passed', 'weight': 3.0}, {'id': 'X2', 'status': 'passed', 'weight': 5.2}, {'id': 'X3', 'status': 'passed', 'weight': 2.5}]
>>> result = inspect_parts(test_parts)
>>> expected_alert_message = 'Overweight part detected!'
>>> alert_message = result['alert_message']
>>> assert alert_message == expected_alert_message
>>> if alert_message == expected_alert_message:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(alert_message))
```
failure_message: Failed: The function does not handle overweight parts correctly.
success_message: Success: The function correctly detects overweight parts!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-quality-control-4'
>>> max_score = 2.0
>>> score = 0
>>> test_parts = [{'id': 'Y1', 'status': 'passed', 'weight': 2.0}, {'id': 'Y2', 'status': 'rejected', 'weight': 3.0}, {'id': 'Y3', 'status': 'passed', 'weight': 1.5}]
>>> result = inspect_parts(test_parts)
>>> expected_approved_count = 2
>>> expected_accepted_parts = ['Y1', 'Y3']
>>> assert result['approved_count'] == expected_approved_count
>>> assert result['accepted_parts'] == expected_accepted_parts
>>> if result['approved_count'] == expected_approved_count and result['accepted_parts'] == expected_accepted_parts:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not correctly skip rejected parts.
success_message: Success: The function correctly skips rejected parts!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-quality-control-5'
>>> max_score = 2.0
>>> score = 0
>>> test_parts = [{'id': 'X1', 'status': 'passed', 'weight': 6.0}, {'id': 'X2', 'status': 'passed', 'weight': 5.2}, {'id': 'X3', 'status': 'passed', 'weight': 2.5}]
>>> result = inspect_parts(test_parts)
>>> expected_alert_message = 'Overweight part detected!'
>>> alert_message = result['alert_message']
>>> assert alert_message == expected_alert_message
>>> if alert_message == expected_alert_message:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(alert_message))
```
failure_message: Failed: The function does not handle multiple overweight parts correctly.
success_message: Success: The function correctly detects multiple overweight parts!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-quality-control-6'
>>> max_score = 3.0
>>> score = 0
>>> test_parts = [{'id': 'X1', 'status': 'rejected', 'weight': 1.0}, {'id': 'X2', 'status': 'rejected', 'weight': 5.2}, {'id': 'X3', 'status': 'passed', 'weight': 2.5}]
>>> result = inspect_parts(test_parts)
>>> expected_alert_message = 'Overweight part detected!'
>>> alert_message = result['alert_message']
>>> assert alert_message == expected_alert_message
>>> if alert_message == expected_alert_message:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(alert_message))
```
failure_message: Failed: The function does not handle rejected overweight parts.
success_message: Success: The function deals with rejected overweight parts!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-dictionaries-quality-control-7'
>>> max_score = 3.0
>>> score = 0
>>> test_parts = [{'id': 'X1', 'status': 'passed', 'weight': 1.8}, {'id': 'X2', 'status': 'rejected', 'weight': 3.2}, {'id': 'X3', 'status': 'passed', 'weight': 2.5}, {'id': 'X4', 'status': 'passed', 'weight': 4.6}, {'id': 'X5', 'status': 'passed', 'weight': 1.9}]
>>> expected_ids = ['X1', 'X3']
>>> qc_result_new = inspect_parts(test_parts)
>>> actual_ids = qc_result_new['accepted_parts']
>>> assert actual_ids == expected_ids, f'Expected `accepted_parts` to be {expected_ids}, but got {actual_ids}.'
>>> if actual_ids == expected_ids:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(expected_ids))
>>> responses = update_responses(question_id, str(actual_ids))
```
failure_message: Failed: The `accepted_parts` list does not contain the expected IDs for a different input.
success_message: Success: The `accepted_parts` list correctly identifies passed parts!


Tests for loops-and-errors-debugging-thermodynamics (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(13.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(55.0)
>>> log_variable('total-points', f'week6-midterm, 1_midterm_exam_q', 55.0)
>>> question_id = 'loops-and-errors-debugging-thermodynamics-1'
>>> max_score = 2.0
>>> score = 0
>>> function_exists = callable(convert_kelvin_to_celsius)
>>> assert function_exists, 'The function `convert_kelvin_to_celsius` must be defined and callable.'
>>> if (function_exists, 'The function `convert_kelvin_to_celsius` must be defined and callable.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(function_exists))
```
failure_message: Failed: The function `convert_kelvin_to_celsius` is not defined correctly.
success_message: Success: The function `convert_kelvin_to_celsius` exists and is callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> question_id = 'loops-and-errors-debugging-thermodynamics-2'
>>> max_score = 5.0
>>> score = 0
>>> exec(base64.b64decode('Y29uZGl0aW9uMSA9IG5wLmlzY2xvc2UoY29udmVydF9rZWx2aW5fdG9fY2Vsc2l1cyhbMjgwLCAyNjAsIDMwMF0pLCBbNi44NSwgMjYuODVdKS5hbGwoKQ==').decode())
>>> exec(base64.b64decode('Y29uZGl0aW9uMiA9IGNvbnZlcnRfa2VsdmluX3RvX2NlbHNpdXMoWzI2MCwgMjUwLCAyNzBdKSA9PSBbXQ==').decode())
>>> exec(base64.b64decode('Y29uZGl0aW9uMyA9IG5wLmlzY2xvc2UoY29udmVydF9rZWx2aW5fdG9fY2Vsc2l1cyhbMjczLjE1LCAzMTBdKSwgWzAuMCwgMzYuODVdKS5hbGwoKQ==').decode())
>>> assert condition1
>>> assert condition2
>>> assert condition3
>>> if condition1 and condition2 and condition3:
...     score = 5.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result))
```
failure_message: Failed: The function does not correctly filter and convert valid Kelvin temperatures.
success_message: Success: The function correctly filters and converts valid Kelvin temperatures!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'loops-and-errors-debugging-thermodynamics-3'
>>> max_score = 2.0
>>> score = 0
>>> exec(base64.b64decode('Y29uZGl0aW9uID0gY29udmVydF9rZWx2aW5fdG9fY2Vsc2l1cyhbXSkgPT0gW10=').decode())
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result))
```
failure_message: Failed: The function does not handle an empty list correctly.
success_message: Success: The function correctly handles an empty input list!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> import numpy as np
>>> question_id = 'loops-and-errors-debugging-thermodynamics-4'
>>> max_score = 4.0
>>> score = 0
>>> exec(base64.b64decode('Y29uZGl0aW9uMSA9IG5wLmlzY2xvc2UoY29udmVydF9rZWx2aW5fdG9fY2Vsc2l1cyhbMzEwLCAyNzUsIDI5MF0pLCBbMzYuODUsIDEuODUsIDE2Ljg1XSkuYWxsKCk=').decode())
>>> exec(base64.b64decode('Y29uZGl0aW9uMiA9IG5wLmlzY2xvc2UoY29udmVydF9rZWx2aW5fdG9fY2Vsc2l1cyhbMjkwLCAyNzMuMTUsIDI4MF0pLCBbMTYuODUsIDAuMCwgNi44NV0pLmFsbCgp').decode())
>>> assert condition1
>>> assert condition2
>>> if condition1 and condition2:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(result))
```
failure_message: Failed: The function does not preserve order of elements.
success_message: Success: The function preserves order of elements!


Tests for numpy-multiple-of-7-check (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(16.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(55.0)
>>> log_variable('total-points', f'week6-midterm, 1_midterm_exam_q', 55.0)
>>> question_id = 'numpy-multiple-of-7-check-1'
>>> max_score = 2.0
>>> score = 0
>>> function_exists = callable(find_multiples_of_7)
>>> assert function_exists, 'The function `find_multiples_of_7` must be defined and callable.'
>>> if (function_exists, 'The function `find_multiples_of_7` must be defined and callable.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(function_exists))
```
failure_message: Failed: The function `find_multiples_of_7` is not defined correctly.
success_message: Success: The function `find_multiples_of_7` exists and is callable!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-multiple-of-7-check-2'
>>> max_score = 3.0
>>> score = 0
>>> docstring_exists = find_multiples_of_7.__doc__ is not None
>>> docstring_word_count = len(find_multiples_of_7.__doc__.split()) if docstring_exists else 0
>>> assert docstring_exists and docstring_word_count >= 8, f'The function `find_multiples_of_7` must have a docstring with at least 8 words. Current docstring word count: {docstring_word_count}'
>>> if (docstring_exists and docstring_word_count >= 8, f'The function `find_multiples_of_7` must have a docstring with at least 8 words. Current docstring word count: {docstring_word_count}'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(docstring_exists))
>>> responses = update_responses(question_id, str(docstring_word_count))
```
failure_message: Failed: The function `find_multiples_of_7` does not have a docstring with at least 8 words.
success_message: Success: The function `find_multiples_of_7` has a docstring with at least 8 words.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-multiple-of-7-check-3'
>>> max_score = 3.0
>>> score = 0
>>> test_input = np.array([7, 14, 21, 28, 35])
>>> expected_output = np.array([True, True, True, True, True])
>>> multiples_test_results = find_multiples_of_7(test_input)
>>> assert np.array_equal(multiples_test_results, expected_output), f'Expected {expected_output}, but got {multiples_test_results}.'
>>> if (np.array_equal(multiples_test_results, expected_output), f'Expected {expected_output}, but got {multiples_test_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(multiples_test_results))
```
failure_message: Failed: The function does not correctly identify positive multiples of 7.
success_message: Success: The function `find_multiples_of_7` correctly identifies positive multiples of 7!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-multiple-of-7-check-4'
>>> max_score = 3.0
>>> score = 0
>>> test_array = np.array([6, 8, 15, 20, 30, 50])
>>> multiples_results = find_multiples_of_7(test_array)
>>> expected_results = np.array([False, False, False, False, False, False])
>>> assert np.array_equal(multiples_results, expected_results), f'Expected {expected_results}, but got {multiples_results}'
>>> if (np.array_equal(multiples_results, expected_results), f'Expected {expected_results}, but got {multiples_results}'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(test_array))
>>> responses = update_responses(question_id, str(multiples_results))
```
failure_message: Failed: The function does not correctly filter out non-multiples of 7.
success_message: Success: The function correctly handles numbers that are NOT multiples of 7!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-multiple-of-7-check-5'
>>> max_score = 2.0
>>> score = 0
>>> test_input = np.array([-7, -14, -21, 7, 14, 21])
>>> expected_output = np.array([False, False, False, True, True, True])
>>> negative_test_results = find_multiples_of_7(test_input)
>>> assert np.array_equal(negative_test_results, expected_output), f'Expected {expected_output}, but got {negative_test_results}.'
>>> if (np.array_equal(negative_test_results, expected_output), f'Expected {expected_output}, but got {negative_test_results}.'):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(negative_test_results))
```
failure_message: Failed: The function does not correctly exclude negative numbers.
success_message: Success: The function correctly handles negative numbers!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'numpy-multiple-of-7-check-6'
>>> max_score = 3.0
>>> score = 0
>>> test_input = np.array([[7, 14, 3], [21, 6, 49]])
>>> expected_shape = test_input.shape
>>> shape_test_results = find_multiples_of_7(test_input).shape
>>> assert shape_test_results == expected_shape, f'Expected shape {expected_shape}, but got {shape_test_results}.'
>>> if (shape_test_results == expected_shape, f'Expected shape {expected_shape}, but got {shape_test_results}.'):
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(shape_test_results))
```
failure_message: Failed: The function does not preserve the shape of the input array.
success_message: Success: The function preserves NumPy array shape!


Tests for remove-duplicates-sorted-array (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(10.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(55.0)
>>> log_variable('total-points', f'week6-midterm, 1_midterm_exam_q', 55.0)
>>> question_id = 'remove-duplicates-sorted-array-1'
>>> max_score = 10.0
>>> score = 0
>>> test_cases = [([1, 1, 2, 2, 3, 4, 4], 4, [1, 2, 3, 4]), ([1, 2, 3, 4, 5], 5, [1, 2, 3, 4, 5]), ([7, 7, 7, 7, 7], 1, [7]), ([0, 0, 1, 1, 2, 3, 3, 4, 5, 5], 6, [0, 1, 2, 3, 4, 5]), ([5], 1, [5]), ([], 0, []), ([10, 10, 20, 30, 30, 40, 50, 50, 60, 70, 70], 7, [10, 20, 30, 40, 50, 60, 70])]
>>> actual_results = []
>>> for nums, expected_k, expected_nums in test_cases:
...     nums_copy = nums[:]
...     result_k = remove_duplicates(nums_copy)
...     actual_results.append((result_k, nums_copy[:result_k]))
...     assert result_k == expected_k, f'Expected {expected_k}, but got {result_k} for input {nums}'
...     assert nums_copy[:result_k] == expected_nums, f'Expected {expected_nums}, but got {nums_copy[:result_k]} for input {nums}'
>>> if result_k == expected_k and nums_copy[:result_k] == expected_nums:
...     score = 10.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_midterm_exam_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not pass all test cases.
success_message: Success: The function correctly handles all cases! Great JOB!
