```python
from pykubegrader.tokens.validate_token import validate_token

validate_token("type the key provided by your instructor here", assignment="week8-quiz")
```

```python
validate_token(assignment="week8-quiz")

# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment(
    "1_quiz_q", "week_8", "quiz", assignment_points=19.0, assignment_tag="week8-quiz"
)

# Initialize Otter
import otter

grader = otter.Notebook("1_quiz_q.ipynb")
```

# â“Quiz: Implicit Equation Plotting

$$
eq1 => y=\frac{2}{1}x+16\ \left\{-11\le x\le-9\right\}
$$

$$
eq2 => y=\frac{1}{2\ }\left|x+7\right|-8\ \left\{-11\le x\le-5\right\}
$$

$$
eq3 => y=-1x+\left(-12\right)\ \left\{-7\le x\le-5\right\}
$$

$$
eq4 => y=\frac{1}{2}x+1.5\ \left\{-7\le x\le-5\right\}
$$

$$
eq5 => y-\left(-4\right)=\frac{1}{2}\left(x-\left(-5\right)\right)\ \left\{-7\le x\le-5\right\}
$$

$$
eq6 => y-\left(-4\right)=-3\left(x-\left(-5\right)\right)\left\{-5\le x\le-4\right\}
$$

$$
eq7 => y-\left(-7\right)=\frac{2}{1}\left(x-\left(-4\right)\right)\left\{-4\le x\le-3\right\}
$$

$$
eq8 => x=-3\left\{-5\le y\le-3\right\}
$$

$$
eq9 => y-\left(-3\right)=\frac{1}{2}\left(x-\left(-3\right)\right)\ \left\{-3\le x\le-1\right\}
$$

$$
eq10 => y-\left(-2\right)=-1\left(x-\left(-1\right)\right)\ \left\{-1\le x\le0\right\}
$$

$$
eq11 => y-\left(-3\right)=1\left(x-0\right)\ \left\{-1\le x\le0\right\}
$$

$$
eq12 => y-\left(-4\right)=\frac{1}{2}\left(x-\left(-1\right)\right)\ \left\{-1\le x\le3\right\}
$$

$$
eq13 => x=3\left\{-2\le y\le-1\right\}
$$

$$
eq14 => y=1\left|x-4\right|-2\ \left\{3\le x\le7\right\}
$$

$$
eq15 => y-1=0\left(x-7\right)\ \left\{6\le x\le7\right\}
$$

$$
eq16 => y-1=1\left(x-6\right)\ \left\{6\le x\le10\right\}
$$

$$
eq17 => y-5=0\left(x-10\right)\ \left\{3\le x\le10\right\}
$$

$$
eq18 => y-5=3\left(x-3\right)\ \left\{3\le x\le4\right\}
$$

$$
eq19 => y-8=\frac{1}{3}\left(x-4\right)\ \left\{-2\le x\le4\right\}
$$

$$
eq20 => y-6=3\left(x-\left(-2\right)\right)\ \left\{-2\le x\le-1\right\}
$$

$$
eq21 => y-9=1\left(x-\left(-1\right)\right)\ \left\{-7\le x\le-1\right\}
$$

$$
eq22 => x=-7\ \left\{3\le y\le5\right\}
$$

$$
eq23 => y-5=2\left(x-\left(-7\right)\right)\ \left\{-11\le x\le-7\right\}
$$

$$
eq24 => y-\left(-3\right)=\frac{1}{2}\left(x-\left(-11\right)\right)\ \left\{-11\le x\le-9\right\}
$$

```python
# Import the necessary libraries
# numpy, matplotlib.pyplot using common alias plt and np
...
```

```python
grader.check("q1-import-libraries")
```

```python
## ðŸ“Œ Some Predefined Mathematical Equations
```

```python
# Predefined functions
def eq1(x):
    return (2 / 1) * x + 16


def eq4(x):
    return (1 / 2) * x + 1.5


def eq5(x):
    return (-4) + (1 / 2) * (x + 5)


def eq6(x):
    return (-4) + (-3) * (x + 5)


def eq7(x):
    return (-7) + (2 / 1) * (x + 4)


def eq8(y):
    return np.full_like(y, -3)


def eq9(x):
    return (-3) + (1 / 2) * (x + 3)


def eq10(x):
    return (-2) - (x + 1)


def eq11(x):
    return (-3) + 1 * (x - 0)


def eq12(x):
    return (-4) + (1 / 2) * (x + 1)


def eq13(y):
    return np.full_like(y, 3)


def eq14(x):
    return 1 * np.abs(x - 4) - 2


def eq15(x):
    return np.full_like(x, 1)


def eq16(x):
    return 1 + 1 * (x - 6)


def eq17(x):
    return np.full_like(x, 5)


def eq18(x):
    return 5 + 3 * (x - 3)


def eq19(x):
    return 8 + (1 / 3) * (x - 4)


def eq20(x):
    return 6 + 3 * (x + 2)


def eq21(x):
    return 9 + 1 * (x + 1)


def eq22(y):
    return np.full_like(y, -3)


def eq23(x):
    return 5 + 2 * (x + 7)


def eq24(x):
    return (-3) + (1 / 2) * (x + 11)
```

```python
# Define EQ2 and EQ3 functions takes x as input and returns the output
# eq2(x) = (1/2) * |x + 7| - 8, module of x + 7 can be calculated using np.abs()
...
# eq3(x) = -1 * x - 12
...
```

```python
grader.check("q2-compute-the-first-equation")
```

```python
# eq 8. y = -7 use np.full_like() to create an array of -7s, like np.full_like(y, -7)
...

# eq 12. y = 3 use np.full_like() to create an array of 3s, like np.full_like(y, 3)
...
```

```python
grader.check("q3-compute-the-next-equation")
```

```python
# Predefined Ranges
```

```python
ranges_x = [
    np.linspace(-11, -9, 100),
    np.linspace(-11, -5, 100),
    np.linspace(-7, -5, 100),
    np.linspace(-7, -5, 100),
    np.linspace(-7, -5, 100),
    np.linspace(-5, -4, 100),
    np.linspace(-4, -3, 100),
    np.linspace(-3, -1, 100),
    np.linspace(-1, 0, 100),
    np.linspace(-1, 0, 100),
    np.linspace(-1, 3, 100),
    np.linspace(3, 7, 100),
    np.linspace(6, 7, 100),
    np.linspace(6, 10, 100),
    np.linspace(3, 10, 100),
    np.linspace(3, 4, 100),
    np.linspace(-2, 4, 100),
    np.linspace(-2, -1, 100),
    np.linspace(-7, -1, 100),
    np.linspace(-11, -7, 100),
    np.linspace(-11, -9, 100),
]
# Functions corresponding to ranges
functions_x = [
    eq1,
    eq2,
    eq3,
    eq4,
    eq5,
    eq6,
    eq7,
    eq9,
    eq10,
    eq11,
    eq12,
    eq14,
    eq15,
    eq16,
    eq17,
    eq18,
    eq19,
    eq20,
    eq21,
    eq23,
    eq24,
]
```

```python
# define ranges_y for the vertical lines using np.linspace() function with 100 points between the given range as np.linspace(start, end, 100)
# eq 8. y = -3, use np.linspace(start, end, 100) to create an array of -3s, like np.linspace(3, 5, 100)
# eq 13. y = 3, use np.linspace(start, end, 100) to create an array of 3s, like np.linspace(-2, -1, 100)
# eq 22. y = -5, use np.linspace(start, end, 100) to create an array of -5s, like np.linspace(-3, -5, 100)
...

# make a list of eq8, eq13, eq22, save to the variable functions_y
...

# Initialize the plot with figure size 10, 6 using plt.subplots() function, save to the variable fig, and ax for the axes, figsize is the optional argument.
...


# Plot x-based functions
for func, x_range in zip(functions_x, ranges_x):
    ax.plot(x_range, func(x_range), "g", linewidth=2)


# Plot the equations y-based functions (vertical lines) using plt.plot() with blue color 'b' and linewidth 2
# using a for loop to iterate over the functions_y and ranges_y using zip() function to combine them
# look at the previous example for the reference of the synax -- you can basicaly copy the syntax and change the variables
...


# Set up plot appearance
# Set the x-axis label using ax.set_xlabel() function - "X-Axis"
# Set the y-axis label using ax.set_ylabel() function - "Y-Axis"
# Set the title using ax.set_title() function - "Graph of Multiple Equations making Eagles Logo"
# Set the grid using ax.grid() function
# Show the plot using plt.show() function
# Set axhline and axvline using plt.axhline() and plt.axvline() functions to draw horizontal and vertical lines
# Set the color of the lines as gray using color='gray'
# Set the linestyle as dashed using linestyle='dashed'
# Set the linewidth as 1 using linewidth=1
...
```

```python
grader.check("q4-compute-the-range-and-plots")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token

validate_token(assignment="week8-quiz")


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week8-quiz", "1_quiz_q")
```

Tests for q1-import-libraries (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> import sys
>>> imported_modules = set(sys.modules.keys())
>>> max_question_points = str(2.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(19.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> log_variable('total-points', f'week8-quiz, 1_quiz_q', 19.0)
>>> question_id = 'q1-import-libraries-1'
>>> max_score = 2.0
>>> score = 0
>>> required_libraries = {'matplotlib.pyplot', 'numpy', 'math'}
>>> missing_libraries = required_libraries - imported_modules
>>> assert not missing_libraries, f'Missing imports: {missing_libraries}. Ensure all required libraries are imported.'
>>> if not missing_libraries:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(imported_modules))
```
failure_message: Failed: One or more required libraries or typing imports are missing.
success_message: Success: All required libraries and typing imports are correctly imported!


Tests for q2-compute-the-first-equation (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(19.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> log_variable('total-points', f'week8-quiz, 1_quiz_q', 19.0)
>>> question_id = 'q2-compute-the-first-equation-1'
>>> max_score = 2.0
>>> score = 0
>>> eq2_edge_cases = [eq2(x) for x in [-11, -7, -5]]
>>> condition = all((isinstance(val, float) for val in eq2_edge_cases))
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(eq2_edge_cases))
```
failure_message: Failed: `eq2()` does not handle absolute values correctly.
success_message: Success: `eq2()` handles absolute values correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> question_id = 'q2-compute-the-first-equation-2'
>>> max_score = 2.0
>>> score = 0
>>> eq3_edge_cases = [eq3(x) for x in [-20, -15, -50]]
>>> condition = all((isinstance(val, int) for val in eq3_edge_cases))
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(eq3_edge_cases))
```
failure_message: Failed: `eq3()` does not handle negative slope correctly.
success_message: Success: `eq3()` handles negative slope correctly!


Tests for q3-compute-the-next-equation (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> max_question_points = str(4.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(19.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> log_variable('total-points', f'week8-quiz, 1_quiz_q', 19.0)
>>> question_id = 'q3-compute-the-next-equation-1'
>>> max_score = 2.0
>>> score = 0
>>> eq8_edge_cases = [eq8(y) for y in [-5, -4, -3]]
>>> condition = all((val == -7 for val in eq8_edge_cases))
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(eq8_edge_cases))
```
failure_message: Failed: `eq8()` does not return constant x values correctly.
success_message: Success: `eq8()` correctly returns constant x values!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> question_id = 'q3-compute-the-next-equation-2'
>>> max_score = 2.0
>>> score = 0
>>> eq13_edge_cases = [eq13(y) for y in [-2, -1]]
>>> condition = all((val == 3 for val in eq13_edge_cases))
>>> assert condition
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(eq13_edge_cases))
```
failure_message: Failed: `eq13()` does not return constant x values correctly.
success_message: Success: `eq13()` correctly returns constant x values!


Tests for q4-compute-the-range-and-plots (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> max_question_points = str(9.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(19.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> log_variable('total-points', f'week8-quiz, 1_quiz_q', 19.0)
>>> question_id = 'q4-compute-the-range-and-plots-1'
>>> max_score = 1.0
>>> score = 0
>>> ranges_y_correct = len(ranges_y) == 3 and all((len(r) == 100 for r in ranges_y)) and np.allclose(ranges_y[0], np.linspace(3, 5, 100)) and np.allclose(ranges_y[1], np.linspace(-2, -1, 100)) and np.allclose(ranges_y[2], np.linspace(-3, -5, 100))
>>> assert ranges_y_correct, 'Expected `ranges_y` to be initialized with correct ranges and 100 points.'
>>> if ranges_y_correct:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(ranges_y_correct))
```
failure_message: Failed: `ranges_y` is not correctly initialized with 100 points.
success_message: Success: `ranges_y` is correctly initialized with 100 points!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> question_id = 'q4-compute-the-range-and-plots-2'
>>> max_score = 1.0
>>> score = 0
>>> functions_y_correct = functions_y == [eq8, eq13, eq22]
>>> assert functions_y_correct, 'Expected `functions_y` to contain [eq8, eq13, eq22].'
>>> if functions_y_correct:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(functions_y_correct))
```
failure_message: Failed: `functions_y` is not correctly initialized with the expected functions.
success_message: Success: `functions_y` is correctly initialized with corresponding functions!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> import matplotlib.pyplot as plt
>>> question_id = 'q4-compute-the-range-and-plots-3'
>>> max_score = 2.0
>>> score = 0
>>> figure_initialized = all(np.isclose(fig.get_size_inches(), (10, 6)))
>>> assert figure_initialized, 'Expected figure size to be (10, 6).'
>>> if figure_initialized:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(figure_initialized))
```
failure_message: Failed: Figure not initialized with the correct size.
success_message: Success: Figure initialized with correct size!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> question_id = 'q4-compute-the-range-and-plots-4'
>>> max_score = 2.0
>>> score = 0
>>> vertical_lines_correct = len(ax.lines) == 26
>>> assert vertical_lines_correct, 'Expected vertical lines to be plotted for each function in `functions_y`.'
>>> if vertical_lines_correct:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(vertical_lines_correct))
```
failure_message: Failed: Vertical lines are not plotted correctly.
success_message: Success: Vertical lines are plotted with the correct functions and ranges!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week8-quiz')
>>> question_id = 'q4-compute-the-range-and-plots-5'
>>> max_score = 3.0
>>> score = 0
>>> labels_title_correct = ax.get_xlabel() == 'X-Axis' and ax.get_ylabel() == 'Y-Axis' and (ax.get_title() == 'Graph of Multiple Equations making Eagles Logo')
>>> assert labels_title_correct, 'Expected correct labels, title, and grid setup.'
>>> if labels_title_correct:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(labels_title_correct))
```
failure_message: Failed: Labels, title, or grid are not set up correctly.
success_message: Success: Labels, title, and grid are set up correctly!
