```python
from pykubegrader.tokens.validate_token import validate_token
validate_token('type the key provided by your instructor here') 

```

```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_quiz_q", "week_4", "quiz", assignment_points = 19.0, assignment_tag = 'week4-quiz')

# Initialize Otter
import otter
grader = otter.Notebook("1_quiz_q.ipynb")
```

# ❓ Quiz: Control Structures for Engineering Applications

Test your understanding of conditional statements (`if-else`) and loops in Python by solving problems relevant to engineering applications, such as sensor data processing, control systems, and automation tasks.

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_q import Question3
Question3().show()

```

## Water Turbidity Monitoring System

Water turbidity is a measure of water clarity, often used to ensure safe drinking water or proper water quality in engineering applications. Turbidity is measured in Nephelometric Turbidity Units (NTU). In this task, you will implement a program to process turbidity data collected over time and check for compliance with safety standards.

The program should:

1. Iterate through a list of turbidity measurements (in NTU).
2. Print a warning message if the turbidity exceeds the safe threshold of 5.0 NTU.
3. Calculate the average turbidity of all measurements and display it.
4. Determine if the average turbidity complies with the safe threshold (≤ 5.0 NTU) and display a summary message.

### Implementation Details

1. Define a list of turbidity measurements: `turbidity_readings = [2.5, 3.0, 5.5, 6.2, 4.8, 3.3]`.
2. Set the safe turbidity threshold: `safe_threshold = 5.0`.
3. Loop through all turbidity measurements and:
   - Print a warning message if the turbidity exceeds the safe threshold.
   - The loop should assign the local variable `turbidity` in each iteration of the for loop. 
4. Calculate the average turbidity of all measurements. Save to the variable `average_turbidity`.
   - Print the average turbidity value. We have provided the code for you to calculate the average turbidity. You can use the `sum()` and `len()` functions to calculate the average, or the numpy library if you prefer.
5. Check if the average turbidity complies with the safe threshold:
   - If the average turbidity is less than or equal to the safe threshold, print a message indicating compliance.
   - If the average turbidity exceeds the safe threshold, print a message indicating non-compliance.
   We have provided the print statements for you. 




```python
# Define the turbidity measurements
...

# Define the safe turbidity threshold
...

def turbidity_warning(turbidity, threshold):

    # Iterate through turbidity readings
    ...
        
        # Check if the turbidity exceeds the safe threshold
        ...
            print(f"Warning: Turbidity exceeds safe threshold at {turbidity} NTU!")

    # Calculate the average turbidity
    ...

    # Print the average turbidity
    print(f"Average turbidity: {average_turbidity:.2f} NTU")

    # Check if the average turbidity is safe
    ...
        print("The average turbidity is within the safe range.")
    ...
        print("The average turbidity exceeds the safe threshold!")

```

```python
grader.check("Water-Turbidity")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token
validate_token()


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week4-quiz", "1_quiz_q")
```

Tests for Water-Turbidity (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(10.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(10.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token()
>>> log_variable('total-points', f'Reading-Week-X, 1_quiz_q', 10.0)
>>> question_id = 'Water-Turbidity-1'
>>> max_score = 1.0
>>> score = 0
>>> condition = isinstance(turbidity_readings, list) and len(turbidity_readings) > 0 and all((isinstance(t, (int, float)) for t in turbidity_readings))
>>> assert condition, 'Turbidity readings must be a list of numerical values.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(turbidity_readings))
```
failure_message: Failed: Turbidity measurements are not defined correctly.
success_message: Success: Turbidity measurements are defined correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Water-Turbidity-2'
>>> max_score = 1.0
>>> score = 0
>>> condition = isinstance(safe_threshold, (int, float)) and safe_threshold > 0
>>> assert condition, 'Safe threshold must be a positive numerical value.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(safe_threshold))
```
failure_message: Failed: The safe turbidity threshold is not defined correctly.
success_message: Success: The safe turbidity threshold is defined correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Water-Turbidity-3'
>>> max_score = 1.0
>>> score = 0
>>> iteration_count = 0
>>> for turbidity in turbidity_readings:
...     iteration_count += 1
>>> expected_count = len(turbidity_readings)
>>> condition = iteration_count == expected_count
>>> assert condition, f'Expected iteration over {expected_count} readings, but iterated {iteration_count} times.'
>>> if condition:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(iteration_count))
>>> responses = update_responses(question_id, str(expected_count))
```
failure_message: Failed: The code does not correctly iterate through turbidity readings.
success_message: Success: The code correctly iterates through turbidity readings!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Water-Turbidity-4'
>>> max_score = 2.0
>>> score = 0
>>> high_turbidity = [turbidity for turbidity in turbidity_readings if turbidity > safe_threshold]
>>> expected_high_turbidity = [5.5, 6.2]
>>> condition = high_turbidity == expected_high_turbidity
>>> assert condition, f'Expected warnings for {expected_high_turbidity}, but got warnings for {high_turbidity}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(high_turbidity))
>>> responses = update_responses(question_id, str(expected_high_turbidity))
```
failure_message: Failed: The function does not identify or warn about high turbidity readings.
success_message: Success: The function correctly identifies and warns about high turbidity readings!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'Water-Turbidity-5'
>>> max_score = 2.0
>>> score = 0
>>> expected_average = sum(turbidity_readings) / len(turbidity_readings)
>>> average_turbidity = sum(turbidity_readings) / len(turbidity_readings)
>>> condition = abs(average_turbidity - expected_average) < 1e-06
>>> assert condition, f'Expected average: {expected_average}, but got: {average_turbidity}.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(average_turbidity))
>>> responses = update_responses(question_id, str(expected_average))
```
failure_message: Failed: The function does not correctly calculate the average turbidity.
success_message: Success: The function correctly calculates the average turbidity!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import re
>>> from io import StringIO
>>> import sys
>>> question_id = 'Water-Turbidity-6'
>>> max_score = 3.0
>>> score = 0
>>> output = StringIO()
>>> sys.stdout = output
>>> turbidity_readings = [2.5, 3.0, 5.5, 6.2, 4.8, 3.3]
>>> safe_threshold = 5.0
>>> expected_outputs = ['Warning: Turbidity exceeds safe threshold at 5.5 NTU!', 'Warning: Turbidity exceeds safe threshold at 6.2 NTU!', 'Average turbidity: 4.22 NTU', 'The average turbidity is within the safe range.']
>>> turbidity_warning(turbidity_readings, safe_threshold)
>>> sys.stdout = sys.__stdout__
>>> output_value = output.getvalue().strip()
>>> condition = all((re.search(re.escape(expected), output_value, re.IGNORECASE) for expected in expected_outputs))
>>> assert condition, f'Expected outputs: {expected_outputs}, but got: {output_value}'
>>> if condition:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function does not handle all cases correctly for the given input.
success_message: Success: The function handles all cases correctly for the given input!
