```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment
from pykubegrader.token.validate_token import validate_token
validate_token('type the key provided by your TA here')

responses = initialize_assignment("1_quiz", "week_1", "quiz", assignment_points = 16.0, assignment_tag = 'week1-quiz')

# Initialize Otter
import otter
grader = otter.Notebook("1_quiz.ipynb")
```

# Calculate Drexel Basketball Player's Game Stats! ðŸ€

![](./assets/figures/Drexel-Basketball.webp)

You are a Drexel Dragons basketball player tracking their performance over the season. You want to calculate how many points you've scored and how many points you need to achieve your target for the season!

- **Player ID**: Your Drexel ID (the first part of your email address).
- The you have played 15 games, scoring an average of 20 points per game.
- The you have also scored 25 bonus points in overtime games.
- The your season target is to score 400 points.

Using Python, write a program to calculate:
1. The **total points scored** by the player so far.
2. The **remaining points needed** to reach their season target.

## Detailed Instructions
1. **Define Player ID**:
   - Create a variable `drexel_id` and assign it the value of your drexel_id -- this is the first part of your email. For example, if your email is `abc123@drexel.edu`, your `drexel_id` would be `abc123`. Note that the `drexel_id` should be a string, you can use single or double quotes to define it. For example, `drexel_id = 'abc123'` or `drexel_id = "abc123"` are both valid.

2. **Define the Number of Games and Bonus Points**:
   - Create a variable `games_played` and assign it the value `15`.
   - Create a variable `average_points` and assign it the value `20`.
   - Create a variable `bonus_points` and assign it the value `25`.
   - Create a variable `season_target` and assign it the value `400`.

3. **Calculate the Total Points**:
   - Multiply the number of games by the average points and save it to `total_game_points`.
   - Add the bonus points to get the total points scored so far, saved as `total_points`.

4. **Calculate Remaining Points**:
   - Subtract the total points scored from the season target to determine the remaining points needed.

5. **Print the Results**:
   - Use f-strings to display the player ID, total points scored, and remaining points needed. We have provided this code for you so you do not have to write it. 


```python
# Player's ID
player_id = ...

# Variables

# Number of games played
...

# Average points per game
...

# Bonus points from overtime
...

# Season target
...

# Calculate total points
...

# Calculate remaining points
...

# DO NOT MODIFY BELOW THIS LINE
# Print the results
print(f"Player ID: {player_id}")
print(f"Total Points Scored: {total_points}")
print(f"Remaining Points to Reach Target: {remaining_points}")
```

```python
grader.check("quiz-1-Drexel-basketball")
```

## ðŸŽ‰ Well Done!

Congratulations on completing the Drexel Basketball Player's Game Stats Calculation! You are one step closer to achieving your season target! ðŸ€

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import validate_token
validate_token()


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week1-quiz", "1_quiz")
```

Tests for quiz-1-Drexel-basketball (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(16.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(16.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token()
>>> log_variable('total-points', f'Reading-Week-X, 1_quiz', 16.0)
>>> question_id = 'quiz-1-Drexel-basketball-1'
>>> max_score = 2.0
>>> score = 0
>>> assert games_played == 15
>>> if games_played == 15:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: games_played not correctly assigned
success_message: games_played correctly assigned
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'quiz-1-Drexel-basketball-2'
>>> max_score = 2.0
>>> score = 0
>>> assert average_points == 20
>>> if average_points == 20:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: average_points not correctly assigned
success_message: average_points correctly assigned
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'quiz-1-Drexel-basketball-3'
>>> max_score = 2.0
>>> score = 0
>>> assert bonus_points == 25
>>> if bonus_points == 25:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: bonus_points not correctly assigned
success_message: bonus_points correctly assigned
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'quiz-1-Drexel-basketball-4'
>>> max_score = 2.0
>>> score = 0
>>> assert season_target == 400
>>> if season_target == 400:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: season_target not correctly assigned
success_message: season_target correctly assigned
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'quiz-1-Drexel-basketball-5'
>>> max_score = 4.0
>>> score = 0
>>> assert total_points == 325
>>> if total_points == 325:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: total_points not correctly calculated
success_message: total_points correctly calculated
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'quiz-1-Drexel-basketball-6'
>>> max_score = 4.0
>>> score = 0
>>> assert remaining_points == 75
>>> if remaining_points == 75:
...     score = 4.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: remaining_points not correctly calculated
success_message: remaining_points correctly calculated
