```python
from pykubegrader.tokens.validate_token import validate_token
validate_token('type the key provided by your instructor here', assignment = 'week7-quiz')

```

```python
validate_token(assignment = 'week7-quiz')

# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_quiz_q", "week_7", "quiz", assignment_points = 28.0, assignment_tag = 'week7-quiz')

# Initialize Otter
import otter
grader = otter.Notebook("1_quiz_q.ipynb")
```

# ❓Quiz: Object-Oriented Programming in Piping Fluid Flow

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_q import Question3
Question3().show()

```

## Question 1 (Points: 12.0): Modify a Pipeline Class

Given the `Pipeline` class below, modify it to:
- Add an `__init__()` method that initializes the `name` and `flow_rate` attributes. The `flow_rate` attribute should default to 0.
- Add an `increase_flow(rate)` method that increases the `flow_rate` attribute. 
- Add a `decrease_flow(rate)` method that decreases `flow_rate` but ensures it doesn’t go below zero.
- Implement a `__str__()` method that returns `"Pipeline: <name>, Flow Rate: <flow_rate> L/s"`.

### Example:
```python
pipe = Pipeline("Main", 50)
pipe.increase_flow(20)
print(pipe)  # Output: Pipeline: Main, Flow Rate: 70 L/s
pipe.decrease_flow(80)
print(pipe)  # Output: Pipeline: Main, Flow Rate: 0 L/s
```

```python
...

pipe = Pipeline("Main", 50)
pipe.increase_flow(20)
print(pipe)  # Output: Pipeline: Main, Flow Rate: 70 L/s
pipe.decrease_flow(80)
print(pipe)  # Output: Pipeline: Main, Flow Rate: 0 L/s
```

```python
grader.check("Free-Response-Modify-a-Pipeline-Class")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token
validate_token(assignment = 'week7-quiz')


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week7-quiz", "1_quiz_q")
```

Tests for Free-Response-Modify-a-Pipeline-Class (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week7-quiz')
>>> max_question_points = str(12.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(12.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week7-quiz')
>>> log_variable('total-points', f'week7-quiz, 1_quiz_q', 12.0)
>>> question_id = 'Free-Response-Modify-a-Pipeline-Class-1'
>>> max_score = 2.0
>>> score = 0
>>> item = Pipeline('Main', 50)
>>> pipeline_initialization = item.name == 'Main' and item.flow_rate == 50
>>> assert pipeline_initialization, f"Expected Pipeline('Main', 50), but got name={item.name}, flow_rate={item.flow_rate}."
>>> if pipeline_initialization:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(pipeline_initialization))
```
failure_message: Failed: Pipeline object does not initialize correctly.
success_message: Success: Pipeline object initializes correctly!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week7-quiz')
>>> question_id = 'Free-Response-Modify-a-Pipeline-Class-2'
>>> max_score = 2.0
>>> score = 0
>>> item = Pipeline('Main')
>>> default_flow_rate_check = item.flow_rate == 0
>>> assert default_flow_rate_check, f'Expected default flow_rate to be 0, but got {item.flow_rate}.'
>>> if default_flow_rate_check:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(default_flow_rate_check))
```
failure_message: Failed: Default flow_rate value is incorrect.
success_message: Success: Default stock value is correctly set to 0!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week7-quiz')
>>> question_id = 'Free-Response-Modify-a-Pipeline-Class-3'
>>> max_score = 2.0
>>> score = 0
>>> item = Pipeline('Main', 10)
>>> item.increase_flow(20)
>>> flow_after_addition = item.flow_rate
>>> assert flow_after_addition == 30, f'Expected flow_rate to be 30 after addition, but got {item.flow_rate}.'
>>> if flow_after_addition == 30:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(flow_after_addition))
```
failure_message: Failed: `increase_flow()` does not increase flow_rate correctly.
success_message: Success: `increase_flow()` correctly increases flow_rate!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week7-quiz')
>>> question_id = 'Free-Response-Modify-a-Pipeline-Class-4'
>>> max_score = 2.0
>>> score = 0
>>> item = Pipeline('Main', 50)
>>> item.decrease_flow(30)
>>> flow_after_removal = item.flow_rate
>>> assert flow_after_removal == 20, f'Expected flow_rate to be 20 after removal, but got {item.flow_rate}.'
>>> if flow_after_removal == 20:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(flow_after_removal))
```
failure_message: Failed: `decrease_flow()` does not decrease flow_rate correctly.
success_message: Success: `decrease_flow()` correctly decreases flow_rate!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week7-quiz')
>>> question_id = 'Free-Response-Modify-a-Pipeline-Class-5'
>>> max_score = 2.0
>>> score = 0
>>> item = Pipeline('Main', 40)
>>> item.decrease_flow(100)
>>> flow_after_excessive_removal = item.flow_rate
>>> assert flow_after_excessive_removal == 0, f'Expected flow_rate to be 0 (not negative), but got {item.flow_rate}.'
>>> if flow_after_excessive_removal == 0:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(flow_after_excessive_removal))
```
failure_message: Failed: Flow rate goes below 0.
success_message: Success: Flow rate does not drop below 0!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token(assignment='week7-quiz')
>>> question_id = 'Free-Response-Modify-a-Pipeline-Class-6'
>>> max_score = 2.0
>>> score = 0
>>> item = Pipeline('Main', 15)
>>> string_representation = str(item)
>>> expected_string = 'Pipeline: Main, Flow Rate: 15 L/s'
>>> assert string_representation == expected_string, f"Expected '{expected_string}', but got '{string_representation}'."
>>> if string_representation == expected_string:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(string_representation))
```
failure_message: Failed: `__str__()` does not return the expected format.
success_message: Success: `__str__()` returns the correct format!
