```python
from pykubegrader.tokens.validate_token import validate_token

validate_token("type the key provided by your instructor here")
```

```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment(
    "1_quiz_2_q", "week_2", "quiz", assignment_points=14.5, assignment_tag="week2-quiz"
)

# Initialize Otter
import otter

grader = otter.Notebook("1_quiz_2_q.ipynb")
```

# â“ Python Basics: Data Structures and Types Quiz ðŸ


```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_2_q import Question1

Question1().show()
```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_2_q import Question2

Question2().show()
```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_2_q import Question3

Question3().show()
```

## Engine Performance Analysis ðŸ”§

### Instructions:

Create a Python dictionary to analyze engine test data. Your solution must include:

1. Create a dictionary named `engine_data` with these EXACT keys and values:
   - `model`: String "V6-3000"
   - `displacement`: Float 3.0
   - `test_temperatures`: List [85, 90, 88, 92, 87]  # in Celsius
   - `performance_metrics`: Dictionary containing:
     - `horsepower`: Float 280.5
     - `torque`: Float 350.8  # Nm

2. Calculate these values (store each in its own variable):
   - `avg_temp`: Average of test_temperatures (use `sum` and `len`)
   - `max_temp`: Maximum temperature from test_temperatures (use `max`)
   - `power_density`: Horsepower divided by displacement

```python
# Create the engine data dictionary with exact values
...


# Calculate required values
...

# We have provided the code to print the results for you
# Print results for verification (optional but helpful)
print(f"Average Temperature: {avg_temp:.1f}Â°C")
print(f"Maximum Temperature: {max_temp}Â°C")
print(f"Power Density: {power_density:.1f} hp/L")
```

```python
grader.check("question-Engine-Data")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token

validate_token()


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week2-quiz", "1_quiz_2_q")
```

Tests for question-Engine-Data (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(6.5)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(6.5)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token()
>>> log_variable('total-points', f'Reading-Week-X, 1_quiz_2_q', 6.5)
>>> question_id = 'question-Engine-Data-1'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, isinstance(engine_data, dict))
>>> assert isinstance(engine_data, dict), 'engine_data must be a dictionary'
>>> if isinstance(engine_data, dict):
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Not a dictionary - make sure you create engine_data as a dictionary
success_message: Dictionary type is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Engine-Data-2'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, str(type(engine_data['model'])))
>>> assert isinstance(engine_data['model'], str), 'model must be a string'
>>> if isinstance(engine_data['model'], str):
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The 'model' value must be a string (str)
success_message: Model data type is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Engine-Data-3'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, str(type(engine_data['displacement'])))
>>> assert isinstance(engine_data['displacement'], float), 'displacement must be a float'
>>> if isinstance(engine_data['displacement'], float):
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The 'displacement' value must be a float (not an int)
success_message: Displacement data type is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Engine-Data-4'
>>> max_score = 0.5
>>> score = 0
>>> assert isinstance(engine_data['test_temperatures'], list), 'test_temperatures must be a list'
>>> if isinstance(engine_data['test_temperatures'], list):
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The 'test_temperatures' value must be a list
success_message: Test temperatures data type is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Engine-Data-5'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, str(type(engine_data['performance_metrics'])))
>>> assert isinstance(engine_data['performance_metrics'], dict), 'performance_metrics must be a dictionary'
>>> if isinstance(engine_data['performance_metrics'], dict):
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The 'performance_metrics' value must be a dictionary
success_message: Performance metrics data type is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Engine-Data-6'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, engine_data['model'])
>>> assert engine_data['model'] == 'V6-3000', "model must be exactly 'V6-3000'"
>>> if engine_data['model'] == 'V6-3000':
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The model must be exactly 'V6-3000' (check spelling and capitalization)
success_message: Model value is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Engine-Data-7'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, abs(engine_data['displacement'] - 3.0))
>>> assert abs(engine_data['displacement'] - 3.0) < 0.001, 'displacement must be exactly 3.0'
>>> if abs(engine_data['displacement'] - 3.0) < 0.001:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The displacement must be exactly 3.0 (as a float)
success_message: Displacement value is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Engine-Data-8'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, engine_data['test_temperatures'])
>>> conditional = engine_data['test_temperatures'] == [85, 90, 88, 92, 87]
>>> assert conditional, 'test_temperatures must match exactly'
>>> if conditional:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The test_temperatures list must be exactly [85, 90, 88, 92, 87]
success_message: Test temperatures list is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Engine-Data-9'
>>> max_score = 0.5
>>> score = 0
>>> assert abs(engine_data['performance_metrics']['horsepower'] - 280.5) < 0.001, 'horsepower must be exactly 280.5'
>>> if abs(engine_data['performance_metrics']['horsepower'] - 280.5) < 0.001:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The horsepower must be exactly 280.5 (as a float)
success_message: Horsepower value is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Engine-Data-10'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, abs(engine_data['performance_metrics']['torque']))
>>> assert abs(engine_data['performance_metrics']['torque'] - 350.8) < 0.001, 'torque must be exactly 350.8'
>>> if abs(engine_data['performance_metrics']['torque'] - 350.8) < 0.001:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: The torque must be exactly 350.8 (as a float)
success_message: Torque value is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Engine-Data-11'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, avg_temp)
>>> assert isinstance(avg_temp, float), 'avg_temp must be a float'
>>> assert abs(avg_temp - 88.4) < 0.001, 'avg_temp calculation incorrect - should be 88.4'
>>> if isinstance(avg_temp, float) and abs(avg_temp - 88.4) < 0.001:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: avg_temp must be a float type (use division with sum and len)
success_message: Average temperature data type is correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Engine-Data-12'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, max_temp)
>>> assert isinstance(max_temp, int), 'max_temp must be an integer'
>>> assert max_temp == 92, 'max_temp calculation incorrect - should be 92'
>>> if isinstance(max_temp, int) and max_temp == 92:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: max_temp must be an integer with value 92 (use max() function)
success_message: Maximum temperature data type and value are correct
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> question_id = 'question-Engine-Data-13'
>>> max_score = 0.5
>>> score = 0
>>> responses = update_responses(question_id, power_density)
>>> assert isinstance(power_density, float), 'power_density must be a float'
>>> assert abs(power_density - 93.5) < 0.1, 'power_density calculation incorrect - should be 93.5'
>>> if isinstance(power_density, float) and abs(power_density - 93.5) < 0.1:
...     score = 0.5
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_2_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: power_density must be a float equal to 93.5 (horsepower/displacement)
success_message: Power density calculation is correct
