```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_quiz_q", "week_5", "quiz", assignment_points = 16.0, assignment_tag = 'week5-quiz')

# Initialize Otter
import otter
grader = otter.Notebook("1_quiz_q.ipynb")
```

# ❓ Practice Quiz: Function Concepts in Smart Grid Systems

![](./assets/figures/smart-grid-pquiz.jpg)

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_q import Question2
Question2().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_q import Question3
Question3().show()

```

## Question 1 (Points: 7.0): Smart Grid Power Monitoring System


 The smart grid constantly monitors and processes power readings from different sources, adjusting load distribution dynamically. In this task, you will implement a function that calculates total power consumption and identifies irregularities.

 The program should:

 1. Accept an arbitrary number of power readings (in kilowatts, kW) using `*args`.

 2. Check if any power reading is above `500 kW`, and print a warning if detected.

 3. Calculate the total power consumed and return it.

 4. Allow an optional `alert_threshold` parameter using `**kwargs` that changes the warning limit dynamically.


 ### Implementation Details

 1. Define a function `monitor_power(*args, **kwargs)`.

 2. Loop through all `*args` values and:

    - Print a warning message if any power reading exceeds `500 kW` or the provided `alert_threshold` from `**kwargs`.

 3. Calculate the total power consumed.

 4. Return the total power as output.

```python
# Define the function to monitor power readings
...

    # Extract the alert threshold from the keyword arguments
    # Recall kwargs is a dictionary, and thus we can use the get method to extract the value of the key "alert_threshold"
    ...
    
    # Initialize an empty list to store warnings
    ...
    
    # Iterate over the readings and check if any reading exceeds the alert threshold
    # If a reading exceeds the alert threshold, append a warning to the warnings list
    # we have provided a sample warning message for you to use, look at the call to determine the correct variable to use
    ...
            warnings.append(f"⚠️ Warning: High power usage detected at {reading} kW!")
    
    # Calculate the total power usage
    # We can use the sum function to calculate the total power usage, store in a variable called total_power
    ...
    
    # Return the total power usage and the warnings
    ...
```

```python
grader.check("Smart-Grid-Function-Analysis")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week5-quiz", "1_quiz_q")
```

Tests for Smart-Grid-Function-Analysis (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> max_question_points = str(7.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(7.0)
>>> log_variable('total-points', f'week5-quiz, 1_quiz_q', 7.0)
>>> question_id = 'Smart-Grid-Function-Analysis-1'
>>> max_score = 2.0
>>> score = 0
>>> total_power, _ = monitor_power(200, 150, 100, 50)
>>> assert total_power == 500, f'Expected 500, got {total_power}'
>>> if total_power == 500:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Function does not correctly calculate total power.
success_message: Success: Function correctly calculates total power!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Smart-Grid-Function-Analysis-2'
>>> max_score = 2.0
>>> score = 0
>>> _, warnings = monitor_power(300, 600, 200, 700)
>>> expected_warnings = ['⚠️ Warning: High power usage detected at 600 kW!', '⚠️ Warning: High power usage detected at 700 kW!']
>>> for expected in expected_warnings:
...     assert expected in warnings, f"Expected '{expected}' in warnings but not found."
>>> if expected in warnings:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Function does not correctly detect high power usage.
success_message: Success: Function correctly detects high power usage!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Smart-Grid-Function-Analysis-3'
>>> max_score = 1.0
>>> score = 0
>>> _, warnings = monitor_power(300, 600, 200, 700, alert_threshold=650)
>>> expected_warning = '⚠️ Warning: High power usage detected at 700 kW!'
>>> assert expected_warning in warnings, f"Expected '{expected_warning}', but got {warnings}"
>>> if expected_warning in warnings:
...     score = 1.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Function does not apply custom alert threshold correctly.
success_message: Success: Function correctly applies custom alert threshold!
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import base64
>>> question_id = 'Smart-Grid-Function-Analysis-4'
>>> max_score = 2.0
>>> score = 0
>>> total_power, warnings = monitor_power()
>>> assert total_power == 0, 'Expected total power to be 0 for empty input.'
>>> assert not warnings, 'Expected no warnings for empty input.'
>>> if total_power == 0 and (not warnings):
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: Function does not handle empty input correctly.
success_message: Success: Function correctly handles edge cases (no readings)!
