```python
from pykubegrader.tokens.validate_token import validate_token
validate_token('type the key provided by your instructor here')

```

```python
# You must make sure to run all cells in sequence using shift + enter or you might encounter errors
from pykubegrader.initialize import initialize_assignment

responses = initialize_assignment("1_quiz_q", "week_3", "quiz", assignment_points = 17.0, assignment_tag = 'week3-quiz')

# Initialize Otter
import otter
grader = otter.Notebook("1_quiz_q.ipynb")
```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_q import Question1
Question1().show()

```

```python
# Run this block of code by pressing Shift + Enter to display the question
from questions._1_quiz_q import Question2
Question2().show()

```

## NumPy: Array Creation in the World of Smart Toasters

![](./assets/figures/small-world-toaster.webp)

### Instructions:
You are designing the internal array processing system for a smart toaster that tracks bread slices! Create a NumPy array of integers from 1 to 10 using NumPy's built-in `np.arange()` method. Write a function `create_array()` that:
1. import numpy using the common alias `np`.
2. Uses `np.arange()` to generate the array, from 1 to 10, think about the `start` and `stop` arguments, which ones are inclusive and exclusive?

```python
...

def create_array():
    # Create an array from 1 to 10
    ...
    return array

# Example usage
result = create_array()
print("Generated array:", result)
```

```python
grader.check("smart-toaster-oven")
```

## Submitting Assignment

Please run the following block of code using `shift + enter` to submit your assignment, you should see your score.

```python
from pykubegrader.tokens.validate_token import validate_token
validate_token()


from pykubegrader.submit.submit_assignment import submit_assignment

submit_assignment("week3-quiz", "1_quiz_q")
```

Tests for smart-toaster-oven (points: None)
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> max_question_points = str(10.0)
>>> earned_points = 0
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> os.environ['TOTAL_POINTS_FREE_RESPONSE'] = str(10.0)
>>> from pykubegrader.tokens.validate_token import validate_token
>>> validate_token()
>>> log_variable('total-points', f'Reading-Week-X, 1_quiz_q', 10.0)
>>> question_id = 'smart-toaster-oven-1'
>>> max_score = 2.0
>>> score = 0
>>> condition = 'np' in globals()
>>> assert condition, 'Failed: `np` is not defined globally.'
>>> if condition:
...     score = 2.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: numpy was not imported as np.
success_message: Success: numpy was imported as np.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import numpy as np
>>> question_id = 'smart-toaster-oven-2'
>>> max_score = 5.0
>>> score = 0
>>> expected_array = np.arange(1, 11)
>>> output_array = create_array()
>>> is_array_correct = np.array_equal(output_array, expected_array)
>>> output_type = isinstance(output_array, np.ndarray)
>>> assert is_array_correct, 'Test failed: The generated array is incorrect.'
>>> assert output_type, 'Test failed: Output should be a NumPy array.'
>>> if is_array_correct and output_type:
...     score = 5.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
>>> responses = update_responses(question_id, str(is_array_correct))
>>> responses = update_responses(question_id, str(output_type))
```
failure_message: Failed: The function did not generate the correct array.
success_message: Success: The function did generate the correct array.
```
>>> from pykubegrader.telemetry import ensure_responses, log_variable, score_question, submit_question, telemetry, update_responses
>>> import os
>>> import inspect
>>> import re
>>> question_id = 'smart-toaster-oven-3'
>>> max_score = 3.0
>>> score = 0
>>> code = inspect.getsource(create_array)
>>> arange_used = bool(re.search('np\\.arange\\(', code))
>>> assert arange_used, 'Test failed: np.arange() must be used in the function.'
>>> if arange_used:
...     score = 3.0
>>> earned_points = float(os.environ.get('EARNED_POINTS', 0))
>>> earned_points += score
>>> log_variable('1_quiz_q', f'{score}, {max_score}', question_id)
>>> os.environ['EARNED_POINTS'] = str(earned_points)
```
failure_message: Failed: The function must generate an array using np.arange().
success_message: Success: The function generates a correct NumPy array using np.arange().
